[1633347311] SERVICE ALERT: sistemabap.com.br;url_check;UNKNOWN;SOFT;1;check_http: Invalid hostname/address - sistemabap.com.br/bap
[1633347321] SERVICE ALERT: sistemabap.com.br;url_check;UNKNOWN;SOFT;2;check_http: Invalid hostname/address - sistemabap.com.br/bap
[1633347331] SERVICE ALERT: sistemabap.com.br;url_check;UNKNOWN;HARD;3;check_http: Invalid hostname/address - sistemabap.com.br/bap
[1633347331] HOST ALERT: sistemabap.com.br;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1633347364] HOST ALERT: sistemabap.com.br;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1633347394] HOST ALERT: sistemabap.com.br;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1633347424] HOST ALERT: sistemabap.com.br;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1633347454] HOST ALERT: sistemabap.com.br;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1633347484] HOST ALERT: sistemabap.com.br;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1633347514] HOST ALERT: sistemabap.com.br;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1633347544] HOST ALERT: sistemabap.com.br;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1633347574] HOST ALERT: sistemabap.com.br;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1633347604] HOST ALERT: sistemabap.com.br;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1633347936] SERVICE ALERT: sistemabap.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.496 second response time
[1633348035] SERVICE ALERT: sistemabap.com.br;url_check;UNKNOWN;HARD;1;check_http: Invalid hostname/address - sistemabap.com.br/login
[1633348712] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/opt/nagios/libexec/check_https, ...) failed. errno is 2: No such file or directory
[1633348932] SERVICE ALERT: sistemabap.com.br;url_check;WARNING;HARD;3;Generic error code.
[1633352593] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;HARD;3;Network failure, cannot contact website
[1633352685] SERVICE ALERT: sistemabap.com.br;url_check;UNKNOWN;HARD;3;check_http: Invalid hostname/address - sistemabap.com.br/login
[1633352795] SERVICE ALERT: sistemabap.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.487 second response time
[1633356761] HOST ALERT: sistemabap.com.br;UP;HARD;1;TCP OK - 0.119 second response time on sistemabap.com.br port 80
[1633362933] SERVICE ALERT: rabbitmq.nasajon.com.br;check_rabbitmq;CRITICAL;SOFT;1;RABBITMQ CRITICAL - messages CRITICAL (5595), messages_ready OK (5566) messages_unacknowledged OK (29) consumers OK (0)
[1633362943] SERVICE ALERT: rabbitmq.nasajon.com.br;check_rabbitmq;CRITICAL;SOFT;2;RABBITMQ CRITICAL - messages CRITICAL (5595), messages_ready OK (5566) messages_unacknowledged OK (29) consumers OK (0)
[1633362954] SERVICE ALERT: rabbitmq.nasajon.com.br;check_rabbitmq;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (5688), messages_ready OK (5658) messages_unacknowledged OK (30) consumers OK (0)
[1633363003] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-processar-documento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - messages CRITICAL (5911), messages_ready OK (5880) messages_unacknowledged OK (31) consumers OK (0)
[1633363013] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-processar-documento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - messages CRITICAL (5911), messages_ready OK (5880) messages_unacknowledged OK (31) consumers OK (0)
[1633363024] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-processar-documento;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (5911), messages_ready OK (5880) messages_unacknowledged OK (31) consumers OK (0)
[1633363209] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-consultar-lote;CRITICAL;SOFT;1;RABBITMQ CRITICAL - messages CRITICAL (3295), messages_ready OK (3286) messages_unacknowledged OK (9) consumers OK (0)
[1633363220] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-consultar-lote;CRITICAL;SOFT;2;RABBITMQ CRITICAL - messages CRITICAL (3295), messages_ready OK (3286) messages_unacknowledged OK (9) consumers OK (0)
[1633363231] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-consultar-lote;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (3295), messages_ready OK (3286) messages_unacknowledged OK (9) consumers OK (0)
[1633363332] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-consultar-lote;OK;HARD;3;RABBITMQ OK - messages OK (3295) messages_ready OK (3289) messages_unacknowledged OK (6) consumers OK (0) All queues under the thresholds
[1633363336] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-processar-documento;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (5992), messages_ready OK (5961) messages_unacknowledged OK (31) consumers OK (0)
[1633363356] SERVICE ALERT: rabbitmq.nasajon.com.br;nf-processar-documento;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (6013), messages_ready OK (5984) messages_unacknowledged OK (29) consumers OK (0)
[1633363499] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - messages CRITICAL (6053), messages_ready OK (6021) messages_unacknowledged OK (32) consumers OK (0)
[1633363510] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - messages CRITICAL (6053), messages_ready OK (6021) messages_unacknowledged OK (32) consumers OK (0)
[1633363521] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (6053), messages_ready OK (6021) messages_unacknowledged OK (32) consumers OK (0)
[1633363991] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (5892), messages_ready OK (5821) messages_unacknowledged OK (71) consumers OK (0)
[1633364721] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;OK;HARD;3;RABBITMQ OK - messages OK (4452) messages_ready OK (4304) messages_unacknowledged OK (148) consumers OK (0) All queues under the thresholds
[1633365052] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (4510), messages_ready OK (4353) messages_unacknowledged OK (157) consumers OK (0)
[1633365063] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;OK;SOFT;2;RABBITMQ OK - messages OK (4493) messages_ready OK (4334) messages_unacknowledged OK (159) consumers OK (0) All queues under the thresholds
[1633365714] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (4566), messages_ready OK (4375) messages_unacknowledged OK (191) consumers OK (0)
[1633365724] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (4568), messages_ready OK (4377) messages_unacknowledged OK (191) consumers OK (0)
[1633365735] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (4568), messages_ready OK (4377) messages_unacknowledged OK (191) consumers OK (0)
[1633366026] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;OK;HARD;3;RABBITMQ OK - messages OK (4481) messages_ready OK (4255) messages_unacknowledged OK (226) consumers OK (0) All queues under the thresholds
[1633368509] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633368519] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate 'meutrabalho.app' will expire in 242 days on 2022-06-03 23:59 +0000/UTC.
[1633368559] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633368579] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633368581] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633368589] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;OK;SOFT;3;SSL OK - Certificate 'meutrabalho.app' will expire in 242 days on 2022-06-03 23:59 +0000/UTC.
[1633368592] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.509 second response time
[1633371460] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (1029), messages_ready OK (0) messages_unacknowledged OK (1029) consumers OK (1)
[1633371471] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (1029), messages_ready OK (0) messages_unacknowledged OK (1029) consumers OK (1)
[1633371482] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (1029), messages_ready OK (0) messages_unacknowledged OK (1029) consumers OK (1)
[1633379958] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (4517), messages_ready OK (4517) messages_unacknowledged OK (0) consumers OK (0)
[1633379969] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (4735), messages_ready OK (4735) messages_unacknowledged OK (0) consumers OK (0)
[1633379979] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (4835), messages_ready OK (4835) messages_unacknowledged OK (0) consumers OK (0)
[1633380170] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;OK;HARD;3;RABBITMQ OK - messages OK (4454) messages_ready OK (4453) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633382256] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (4247), messages_ready OK (4244) messages_unacknowledged OK (3) consumers OK (0)
[1633382267] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (4268), messages_ready OK (4265) messages_unacknowledged OK (3) consumers OK (0)
[1633382277] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (4306), messages_ready OK (4303) messages_unacknowledged OK (3) consumers OK (0)
[1633382398] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (4525), messages_ready OK (4522) messages_unacknowledged OK (3) consumers OK (0)
[1633388227] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (3513), messages_ready OK (3513) messages_unacknowledged OK (0) consumers OK (0)
[1633388239] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (3516), messages_ready OK (3516) messages_unacknowledged OK (0) consumers OK (0)
[1633388251] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (3521), messages_ready OK (3521) messages_unacknowledged OK (0) consumers OK (0)
[1633391536] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (4508), messages_ready OK (4508) messages_unacknowledged OK (0) consumers OK (0)