[1637052951] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637052976] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5651 bytes in 5.144 second response time
[1637065901] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637065923] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 62630 bytes in 2.302 second response time
[1637067539] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637067575] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1637067611] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1637067737] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 65419 bytes in 6.905 second response time
[1637075449] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (7003), messages_ready OK (7002) messages_unacknowledged OK (1) consumers OK (0)
[1637075471] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (7490), messages_ready OK (7489) messages_unacknowledged OK (1) consumers OK (0)
[1637075491] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (7941), messages_ready OK (7940) messages_unacknowledged OK (1) consumers OK (0)
[1637075612] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (10216), messages_ready OK (10215) messages_unacknowledged OK (1) consumers OK (0)
[1637081252] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (23) messages_ready OK (22) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1637081533] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (7190), messages_ready OK (7190) messages_unacknowledged OK (0) consumers OK (0)
[1637081554] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (7627), messages_ready OK (7627) messages_unacknowledged OK (0) consumers OK (0)
[1637081575] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (7995), messages_ready OK (7995) messages_unacknowledged OK (0) consumers OK (0)
[1637081615] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1637082153] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637082184] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 62426 bytes in 11.096 second response time
[1637082280] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637082300] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 62422 bytes in 0.068 second response time
[1637083214] SERVICE ALERT: conta.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 791 bytes in 0.313 second response time
[1637083234] SERVICE ALERT: conta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 355 bytes in 0.325 second response time
[1637083374] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637083410] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1637083446] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1637083491] SERVICE ALERT: sistemabap.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5651 bytes in 5.004 second response time
[1637088302] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (701), messages_ready OK (0) messages_unacknowledged OK (701) consumers OK (1)
[1637088323] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (708), messages_ready OK (0) messages_unacknowledged OK (708) consumers OK (1)
[1637088344] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (710), messages_ready OK (0) messages_unacknowledged OK (710) consumers OK (1)
[1637088744] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;OK;HARD;3;RABBITMQ OK - messages OK (31) messages_ready OK (0) messages_unacknowledged OK (31) consumers OK (2) All queues under the thresholds
[1637088887] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637088912] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5651 bytes in 5.055 second response time
[1637089925] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (320), messages_ready OK (319) messages_unacknowledged OK (1) consumers OK (0)
[1637089946] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (115), messages_ready OK (114) messages_unacknowledged OK (1) consumers OK (0)
[1637089967] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1637097221] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (15016), messages_ready OK (15016) messages_unacknowledged OK (0) consumers OK (0)
[1637103401] SERVICE ALERT: telemetria.nasajon.com.br;post_api;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1637103421] SERVICE ALERT: telemetria.nasajon.com.br;post_api;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 335 bytes in 0.323 second response time