[1633480370] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633480380] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.321 second response time
[1633480460] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633480491] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633480501] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.525 second response time
[1633483299] SERVICE ALERT: jenkins.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483299] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483299] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483299] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483299] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483299] SERVICE ALERT: admin.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483299] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483300] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483301] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483301] SERVICE ALERT: conta.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483302] SERVICE ALERT: sentry.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483303] SERVICE ALERT: nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483304] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483305] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483306] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483306] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483307] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483309] SERVICE ALERT: admin.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483309] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483309] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483309] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483309] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483310] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483311] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483312] SERVICE ALERT: sentry.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483313] SERVICE ALERT: nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483314] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483315] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483316] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483317] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;2;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483319] SERVICE ALERT: admin.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483319] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483319] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483319] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483319] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483320] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483321] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483322] SERVICE ALERT: sentry.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483323] SERVICE ALERT: nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483324] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483325] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483326] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483327] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;HARD;3;(No output on stdout) stderr: /bin/sh: 1: Syntax error: Unterminated quoted string
[1633483478] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 0.352 second response time
[1633483479] SERVICE ALERT: conta.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 379 bytes in 0.361 second response time
[1633483484] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.360 second response time
[1633483487] SERVICE ALERT: sentry.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 FOUND - 616 bytes in 0.336 second response time
[1633483564] SERVICE ALERT: lp.nasajon.com.br;url_check;UNKNOWN;HARD;3;check_http: Invalid hostname/address - lp.nasajon.com.br/testeservidor
[1633483565] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 12132 bytes in 0.312 second response time
[1633483565] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1223 bytes in 0.342 second response time
[1633483566] SERVICE ALERT: multinotas.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.352 second response time
[1633483568] SERVICE ALERT: jenkins.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1077 bytes in 0.312 second response time
[1633483568] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.375 second response time
[1633483568] SERVICE ALERT: admin.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 845 bytes in 0.384 second response time
[1633483569] SERVICE ALERT: sistemabap.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.502 second response time
[1633483570] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.354 second response time
[1633483570] SERVICE ALERT: nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 191184 bytes in 0.541 second response time
[1633483571] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.336 second response time
[1633483573] SERVICE ALERT: login.nasajon.com.br;url_check;UNKNOWN;HARD;3;check_http: Invalid hostname/address - login.nasajon.com.br/module.php/nasajon/recuperar_senha.php
[1633483573] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.382 second response time
[1633483748] SERVICE ALERT: admin.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.316 second response time
[1633483748] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.315 second response time
[1633483749] SERVICE ALERT: multinotas.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.320 second response time
[1633483749] SERVICE ALERT: projetos.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.346 second response time
[1633483752] SERVICE ALERT: nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.350 second response time
[1633483753] SERVICE ALERT: atendimento.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.322 second response time
[1633483754] SERVICE ALERT: maladireta.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.326 second response time
[1633483755] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.343 second response time
[1633483757] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.340 second response time
[1633483757] SERVICE ALERT: educacional.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.363 second response time
[1633483757] SERVICE ALERT: sistemabap.com.br;url_check;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.528 second response time
[1633483758] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.332 second response time
[1633483758] SERVICE ALERT: admin.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.322 second response time
[1633483759] SERVICE ALERT: multinotas.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.317 second response time
[1633483759] SERVICE ALERT: projetos.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.324 second response time
[1633483762] SERVICE ALERT: nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.312 second response time
[1633483763] SERVICE ALERT: atendimento.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.307 second response time
[1633483764] SERVICE ALERT: maladireta.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.331 second response time
[1633483765] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.312 second response time
[1633483767] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.320 second response time
[1633483768] SERVICE ALERT: sistemabap.com.br;url_check;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.474 second response time
[1633483769] SERVICE ALERT: admin.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.318 second response time
[1633483769] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.333 second response time
[1633483770] SERVICE ALERT: multinotas.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.323 second response time
[1633483770] SERVICE ALERT: projetos.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.315 second response time
[1633483773] SERVICE ALERT: nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.321 second response time
[1633483774] SERVICE ALERT: atendimento.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.310 second response time
[1633483775] SERVICE ALERT: maladireta.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.315 second response time
[1633483776] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.315 second response time
[1633483778] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.314 second response time
[1633483778] SERVICE ALERT: sistemabap.com.br;url_check;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 272 bytes in 0.479 second response time
[1633483965] SERVICE ALERT: login.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.329 second response time
[1633484066] SERVICE ALERT: admin.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 845 bytes in 0.388 second response time
[1633484129] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.345 second response time
[1633484293] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.349 second response time
[1633484295] SERVICE ALERT: jenkins.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 1507288 bytes in 2.189 second response time
[1633484302] SERVICE ALERT: educacional.nasajon.com.br;url_check;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 500 Internal Server Error - 565 bytes in 0.360 second response time
[1633484563] SERVICE ALERT: multinotas.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.345 second response time
[1633484563] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1223 bytes in 0.327 second response time
[1633484563] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 12132 bytes in 0.336 second response time
[1633484564] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.436 second response time
[1633484566] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.419 second response time
[1633484569] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.338 second response time
[1633484573] SERVICE ALERT: nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 191184 bytes in 0.550 second response time
[1633484819] SERVICE ALERT: sistemabap.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.528 second response time
[1633487266] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633487277] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.485 second response time
[1633487834] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 60680 bytes in 0.096 second response time
[1633489222] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (147), messages_ready OK (146) messages_unacknowledged OK (1) consumers OK (0)
[1633489227] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (127), messages_ready OK (126) messages_unacknowledged OK (1) consumers OK (0)
[1633489235] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (306), messages_ready OK (305) messages_unacknowledged OK (1) consumers OK (0)
[1633489238] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (158), messages_ready OK (156) messages_unacknowledged OK (2) consumers OK (0)
[1633489246] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (405), messages_ready OK (404) messages_unacknowledged OK (1) consumers OK (0)
[1633489249] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (689), messages_ready OK (688) messages_unacknowledged OK (1) consumers OK (0)
[1633489279] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (1752), messages_ready OK (1750) messages_unacknowledged OK (2) consumers OK (0)
[1633489287] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (79) messages_ready OK (78) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633489311] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633489322] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.538 second response time
[1633491172] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633491182] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 64486 bytes in 0.103 second response time
[1633491552] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633491562] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 63489 bytes in 0.130 second response time
[1633492021] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633492032] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.519 second response time
[1633492122] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633492133] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.521 second response time
[1633493243] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633493254] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.493 second response time
[1633494425] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633494435] SERVICE ALERT: login.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.339 second response time
[1633495831] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633495851] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633495852] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633495861] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;OK;SOFT;3;SSL OK - Certificate 'meutrabalho.app' will expire in 240 days on 2022-06-03 23:59 +0000/UTC.
[1633495863] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.502 second response time
[1633496499] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (1471), messages_ready OK (1469) messages_unacknowledged OK (2) consumers OK (0)
[1633497251] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (87) messages_ready OK (85) messages_unacknowledged OK (2) consumers OK (0) All queues under the thresholds
[1633498923] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633498933] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.335 second response time
[1633501035] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633501045] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.384 second response time
[1633501653] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633501664] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.506 second response time
[1633503774] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633503785] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.404 second response time
[1633504514] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633504525] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.515 second response time
[1633510442] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633510452] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.382 second response time
[1633514427] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (401), messages_ready OK (400) messages_unacknowledged OK (1) consumers OK (0)
[1633514439] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (364), messages_ready OK (363) messages_unacknowledged OK (1) consumers OK (0)
[1633514449] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (323), messages_ready OK (322) messages_unacknowledged OK (1) consumers OK (0)
[1633514510] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;OK;HARD;3;RABBITMQ OK - messages OK (66) messages_ready OK (65) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633521426] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521436] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.370 second response time
[1633521456] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521476] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633521496] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633521536] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.324 second response time
[1633521596] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521607] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.475 second response time
[1633521627] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521637] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.431 second response time
[1633521657] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521668] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.324 second response time
[1633521689] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521709] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633521729] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633521750] SERVICE FLAPPING ALERT: pontoweb.nasajon.com.br;url_check;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1633521750] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.376 second response time
[1633521780] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521800] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633521820] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633521850] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.327 second response time
[1633521880] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521890] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.347 second response time
[1633521910] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521930] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633521940] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.420 second response time
[1633521970] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633521980] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.334 second response time
[1633522000] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522010] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.326 second response time
[1633522030] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522041] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.378 second response time
[1633522081] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522092] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.802 second response time
[1633522094] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522106] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 1.921 second response time
[1633522122] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522142] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522162] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633522175] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522183] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 1.181 second response time
[1633522189] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 4.621 second response time
[1633522223] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522243] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522254] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.998 second response time
[1633522274] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522275] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522276] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522286] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.424 second response time
[1633522294] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522296] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522306] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.345 second response time
[1633522314] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633522354] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.413 second response time
[1633522374] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522394] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522414] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633522425] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.328 second response time
[1633522445] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522465] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522475] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.335 second response time
[1633522495] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522515] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522535] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633522552] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522554] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cancelamento
[1633522554] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522554] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-lote;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-lote
[1633522555] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522556] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.672 second response time
[1633522559] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522559] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfe
[1633522562] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-por-tenant
[1633522562] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cancelamento
[1633522563] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cte
[1633522563] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador
[1633522563] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-crawler-nfe
[1633522563] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce
[1633522563] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-gaps
[1633522563] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cce
[1633522564] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas
[1633522564] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] comparar-foto-marcacao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/comparar-foto-marcacao
[1633522564] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.325 second response time
[1633522565] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cce;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cce
[1633522565] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.331 second response time
[1633522565] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_enviar;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_enviar
[1633522565] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.318 second response time
[1633522565] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cancelamento
[1633522566] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.308 second response time
[1633522566] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-lote;OK;SOFT;2;RABBITMQ OK - messages OK (245) messages_ready OK (243) messages_unacknowledged OK (2) consumers OK (0) All queues under the thresholds
[1633522566] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-evento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-evento
[1633522566] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.302 second response time
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-manifestacao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-manifestacao
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-gestor;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-gestor
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] envio;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/envio
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-documento-orfao
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-manifestar-nota
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] envio;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/envio
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-tabela-sincronia
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] slas;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/slas
[1633522568] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-trata-erros;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-trata-erros
[1633522569] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador
[1633522569] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfe
[1633522571] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 2.131 second response time
[1633522571] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cancelamento
[1633522572] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-colaborador
[1633522572] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522572] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.315 second response time
[1633522572] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] remover-no-sincronia;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/remover-no-sincronia
[1633522572] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-por-tenant
[1633522572] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador-individual;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador-individual
[1633522573] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-protocolo;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-protocolo
[1633522573] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cancelamento
[1633522573] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar-lote;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar-lote
[1633522573] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_processar
[1633522573] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cte
[1633522574] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador
[1633522575] SERVICE ALERT: admin.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.338 second response time
[1633522575] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-gaps
[1633522575] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.323 second response time
[1633522575] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522575] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-crawler-nfe
[1633522575] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cce
[1633522576] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.320 second response time
[1633522576] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.307 second response time
[1633522576] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (1) All queues under the thresholds
[1633522576] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.324 second response time
[1633522576] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.313 second response time
[1633522576] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] comparar-foto-marcacao;OK;SOFT;2;RABBITMQ OK - messages OK (9) messages_ready OK (0) messages_unacknowledged OK (9) consumers OK (0) All queues under the thresholds
[1633522576] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cce;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522577] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.343 second response time
[1633522577] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_enviar;OK;SOFT;2;RABBITMQ OK - messages OK (1) messages_ready OK (0) messages_unacknowledged OK (1) consumers OK (1) All queues under the thresholds
[1633522577] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;OK;SOFT;3;RABBITMQ OK - messages OK (7) messages_ready OK (0) messages_unacknowledged OK (7) consumers OK (0) All queues under the thresholds
[1633522577] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-evento;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522580] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] envio;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522580] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-manifestacao;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522580] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-gestor;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522580] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;OK;SOFT;2;RABBITMQ OK - messages OK (688) messages_ready OK (688) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-trata-erros;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] envio;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] slas;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (1) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;OK;SOFT;2;RABBITMQ OK - messages OK (1) messages_ready OK (0) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522581] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522582] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cancelamento;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522582] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 503 Service Temporarily Unavailable - 835 bytes in 0.318 second response time
[1633522582] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.353 second response time
[1633522583] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522583] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] remover-no-sincronia;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522583] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522584] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador-individual;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522584] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-protocolo;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522584] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522585] SERVICE ALERT: admin.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 845 bytes in 0.352 second response time
[1633522585] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar-lote;OK;SOFT;2;RABBITMQ OK - messages OK (1199) messages_ready OK (1182) messages_unacknowledged OK (17) consumers OK (0) All queues under the thresholds
[1633522585] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;OK;SOFT;2;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (1) All queues under the thresholds
[1633522585] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522586] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522586] SERVICE ALERT: login.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.387 second response time
[1633522586] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.388 second response time
[1633522586] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.344 second response time
[1633522586] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1223 bytes in 0.332 second response time
[1633522586] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522587] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522587] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633522612] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522617] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522623] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.680 second response time
[1633522637] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522657] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633522703] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522703] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 7.429 second response time
[1633522723] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522734] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 1.875 second response time
[1633522778] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.701 second response time
[1633522798] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522809] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.340 second response time
[1633522839] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522850] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.328 second response time
[1633522870] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522890] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633522910] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633522940] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.406 second response time
[1633522966] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522970] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633522977] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.388 second response time
[1633522981] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.331 second response time
[1633523021] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523041] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523051] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.350 second response time
[1633523071] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523091] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523098] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523110] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.384 second response time
[1633523111] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523126] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523131] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.353 second response time
[1633523136] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.404 second response time
[1633523170] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523171] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523180] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.328 second response time
[1633523182] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.399 second response time
[1633523210] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523212] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523222] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.351 second response time
[1633523230] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523242] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 1.495 second response time
[1633523242] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523262] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523273] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.370 second response time
[1633523292] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523293] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523302] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.335 second response time
[1633523313] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523322] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523323] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.374 second response time
[1633523333] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.335 second response time
[1633523353] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523353] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523363] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.336 second response time
[1633523363] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.330 second response time
[1633523383] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523403] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523413] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523423] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523424] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.427 second response time
[1633523434] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.367 second response time
[1633523444] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523454] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523454] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.688 second response time
[1633523457] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523462] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523464] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.408 second response time
[1633523473] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 5.853 second response time
[1633523482] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523484] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523493] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523502] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523504] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523513] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523514] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523515] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.532 second response time
[1633523533] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523534] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523544] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.505 second response time
[1633523546] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523554] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523566] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523567] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 2.508 second response time
[1633523586] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523606] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.700 second response time
[1633523627] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523645] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 8.077 second response time
[1633523646] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523648] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cancelamento
[1633523650] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] remover-no-sincronia;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/remover-no-sincronia
[1633523650] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-converter;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-converter
[1633523650] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte
[1633523650] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-por-tenant
[1633523650] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-colaborador
[1633523652] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-trata-erros;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-trata-erros
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar-lote;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar-lote
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] criar-no-sincronia;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/criar-no-sincronia
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador-individual;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador-individual
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cte
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_processar
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador-individual;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador-individual
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-protocolo;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-protocolo
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas-violacao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas-violacao
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce-cancelamento
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] campanha-montagem;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/campanha-montagem
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-regras;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-regras
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-schema-desconhecido
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] comparar-foto-marcacao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/comparar-foto-marcacao
[1633523653] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] sincronia-configuracao-atualizar;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/sincronia-configuracao-atualizar
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cce;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cce
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/meurh/calculo-salario
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-gaps
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-lote;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-lote
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-processar-documento
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cancelamento
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_enviar;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_enviar
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-no-sincronia;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-no-sincronia
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-crawler-nfe
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cce
[1633523654] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-evento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-evento
[1633523656] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1223 bytes in 0.384 second response time
[1633523657] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523657] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] envio;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/envio
[1633523657] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-documento-orfao
[1633523659] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-gestor;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-gestor
[1633523660] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador
[1633523660] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-manifestar-nota
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-por-tenant
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-colaborador
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-manifestacao;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-manifestacao
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-tabela-sincronia
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfe
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] envio;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/envio
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] slas;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/slas
[1633523661] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar
[1633523662] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cancelamento;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cancelamento
[1633523663] SERVICE ALERT: admin.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 7.534 second response time
[1633523664] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.309 second response time
[1633523665] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.323 second response time
[1633523665] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.322 second response time
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar-lote;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar-lote
[1633523665] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.337 second response time
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cte
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_processar
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-protocolo;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-protocolo
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas-violacao;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas-violacao
[1633523665] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce-cancelamento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce-cancelamento
[1633523666] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-regras;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-regras
[1633523666] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador-individual;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador-individual
[1633523666] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-schema-desconhecido
[1633523666] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] campanha-montagem;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/campanha-montagem
[1633523666] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.332 second response time
[1633523666] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] sincronia-configuracao-atualizar;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/sincronia-configuracao-atualizar
[1633523666] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] comparar-foto-marcacao;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/comparar-foto-marcacao
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cce;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cce
[1633523667] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.304 second response time
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-gaps
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-lote;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-lote
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/meurh/calculo-salario
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-no-sincronia;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-no-sincronia
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-processar-documento
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-crawler-nfe
[1633523667] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.311 second response time
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cancelamento
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-evento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-evento
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_enviar;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_enviar
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cce
[1633523667] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] envio;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/envio
[1633523668] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-documento-orfao
[1633523670] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-gestor;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-gestor
[1633523671] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador
[1633523671] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-manifestar-nota
[1633523671] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-por-tenant
[1633523672] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-colaborador
[1633523672] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-manifestacao;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-manifestacao
[1633523672] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] envio;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/envio
[1633523672] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfe
[1633523673] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] slas;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/slas
[1633523673] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-tabela-sincronia
[1633523673] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar
[1633523673] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor
[1633523673] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cancelamento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cancelamento
[1633523673] SERVICE ALERT: admin.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 845 bytes in 0.465 second response time
[1633523673] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cancelamento
[1633523674] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.303 second response time
[1633523675] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.315 second response time
[1633523675] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] remover-no-sincronia;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/remover-no-sincronia
[1633523676] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-converter;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-converter
[1633523676] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar-lote;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar-lote
[1633523676] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte
[1633523676] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao
[1633523677] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 502 Bad Gateway - 674 bytes in 0.343 second response time
[1633523677] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 502 Bad Gateway : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cte
[1633523679] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-trata-erros;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-trata-erros
[1633523679] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 2.497 second response time
[1633523679] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 3.634 second response time
[1633523679] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.582 second response time
[1633523680] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] criar-no-sincronia;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/criar-no-sincronia
[1633523680] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador-individual;CRITICAL;SOFT;2;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador-individual
[1633523684] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.391 second response time
[1633523686] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.621 second response time
[1633523687] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523687] SERVICE ALERT: login.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.376 second response time
[1633523688] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-regras;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-regras
[1633523689] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] comparar-foto-marcacao;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/comparar-foto-marcacao
[1633523689] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-no-sincronia;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-no-sincronia
[1633523689] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_enviar;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_enviar
[1633523692] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/emails_processar
[1633523692] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-protocolo;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-protocolo
[1633523692] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas-violacao;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas-violacao
[1633523693] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce-cancelamento;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce-cancelamento
[1633523693] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-manifestacao;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 504 Gateway Time-out : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-manifestacao
[1633523693] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador-individual;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador-individual
[1633523693] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-schema-desconhecido
[1633523693] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] campanha-montagem;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/campanha-montagem
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] sincronia-configuracao-atualizar;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/sincronia-configuracao-atualizar
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-gaps
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cce;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cce
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/meurh/calculo-salario
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfce
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-evento;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-evento
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-crawler-nfe
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/recalculo-por-apuracao-trabalhador
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cce
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-lote;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-consultar-lote
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor-cancelamento
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] envio;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/envio
[1633523694] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-processar-documento
[1633523695] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-documento-orfao
[1633523695] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-gestor;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/notificacao-gestor
[1633523696] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/pontoweb/resumo-trabalhador
[1633523696] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-manifestar-nota
[1633523697] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1223 bytes in 0.452 second response time
[1633523698] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-nfe
[1633523698] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] envio;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/maladireta/envio
[1633523699] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] slas;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/slas
[1633523699] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/atualizar-tabela-sincronia
[1633523699] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523699] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-emissor
[1633523699] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-cte-consultar
[1633523699] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cancelamento;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte-cancelamento
[1633523699] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cancelamento
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-por-tenant;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar-lote;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] comparar-foto-marcacao;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador-individual;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-converter;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-regras;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523700] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-colaborador;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523701] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] remover-no-sincronia;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/diretorio/remover-no-sincronia
[1633523701] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte;CRITICAL;HARD;3;RABBITMQ CRITICAL - Received 500 read timeout : https://rabbitmq.nasajon.com.br:443/api/queues/multinotas/nf-armazenar-cte
[1633523703] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523703] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-no-sincronia;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523703] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_enviar;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523703] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-trata-erros;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523703] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] criar-no-sincronia;OK;SOFT;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523704] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cte;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523705] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] emails_processar;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523706] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-protocolo;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523709] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523709] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-manifestacao;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523709] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce-cancelamento;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523709] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas-violacao;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador-individual;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] campanha-montagem;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;OK;HARD;3;RABBITMQ OK - messages OK (28) messages_ready OK (26) messages_unacknowledged OK (2) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfce;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-gaps;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523710] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cce;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-consultar-lote;OK;HARD;3;RABBITMQ OK - messages OK (110) messages_ready OK (107) messages_unacknowledged OK (3) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] envio;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] sincronia-configuracao-atualizar;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-evento;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;OK;HARD;3;RABBITMQ OK - messages OK (1) messages_ready OK (1) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-crawler-nfe;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 1.193 second response time
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cce;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;OK;HARD;3;RABBITMQ OK - messages OK (3) messages_ready OK (3) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] notificacao-gestor;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-manifestar-nota;OK;HARD;3;RABBITMQ OK - messages OK (1) messages_ready OK (0) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor-cancelamento;OK;HARD;3;RABBITMQ OK - messages OK (7) messages_ready OK (0) messages_unacknowledged OK (7) consumers OK (0) All queues under the thresholds
[1633523711] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-documento-orfao;OK;HARD;3;RABBITMQ OK - messages OK (688) messages_ready OK (688) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523712] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] resumo-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523715] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-cte-consultar;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] slas;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (1) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-nfe;OK;HARD;3;RABBITMQ OK - messages OK (10) messages_ready OK (10) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cancelamento;OK;HARD;3;RABBITMQ OK - messages OK (1) messages_ready OK (1) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-emissor;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[MALA DIRETA] envio;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] atualizar-tabela-sincronia;OK;HARD;3;RABBITMQ OK - messages OK (89) messages_ready OK (89) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte-cancelamento;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523716] SERVICE ALERT: rabbitmq.nasajon.com.br;[DIRETORIO] remover-no-sincronia;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523717] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-armazenar-cte;OK;HARD;3;RABBITMQ OK - messages OK (0) messages_ready OK (0) messages_unacknowledged OK (0) consumers OK (0) All queues under the thresholds
[1633523729] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523749] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523760] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.834 second response time
[1633523822] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523822] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523833] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.691 second response time
[1633523834] SERVICE ALERT: multinotas.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 1.858 second response time
[1633523842] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523853] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523862] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523864] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.867 second response time
[1633523873] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.572 second response time
[1633523895] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523904] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523914] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.331 second response time
[1633523915] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523918] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 4.364 second response time
[1633523934] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523935] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523945] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.681 second response time
[1633523954] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633523974] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633523975] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523984] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633523994] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.471 second response time
[1633523995] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633524015] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633524020] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524024] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524035] SERVICE FLAPPING ALERT: atendimento.nasajon.com.br;url_check;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1633524035] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.415 second response time
[1633524036] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 1.832 second response time
[1633524038] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 7.848 second response time
[1633524055] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524056] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524062] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 7.216 second response time
[1633524068] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524075] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633524076] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633524086] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.396 second response time
[1633524086] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.507 second response time
[1633524088] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633524102] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524106] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524108] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633524114] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 1.482 second response time
[1633524120] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 4.534 second response time
[1633524150] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524161] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.583 second response time
[1633524168] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 9.299 second response time
[1633524211] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524222] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.347 second response time
[1633524245] SERVICE FLAPPING ALERT: pontoweb.nasajon.com.br;url_check;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1633524252] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524272] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633524283] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.331 second response time
[1633524284] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633524294] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.492 second response time
[1633524324] SERVICE FLAPPING ALERT: atendimento.nasajon.com.br;url_check;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1633525453] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525473] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633525484] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.454 second response time
[1633525584] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525595] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.354 second response time
[1633525615] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525625] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.370 second response time
[1633525645] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525654] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525656] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.371 second response time
[1633525672] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 7.630 second response time
[1633525702] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525712] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.344 second response time
[1633525716] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525727] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.359 second response time
[1633525757] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525768] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.328 second response time
[1633525788] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525799] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.756 second response time
[1633525874] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525875] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (515), messages_ready OK (513) messages_unacknowledged OK (2) consumers OK (0)
[1633525885] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.445 second response time
[1633525892] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (482), messages_ready OK (480) messages_unacknowledged OK (2) consumers OK (0)
[1633525912] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (471), messages_ready OK (469) messages_unacknowledged OK (2) consumers OK (0)
[1633525945] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633525961] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 5.956 second response time
[1633526046] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526061] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 5.174 second response time
[1633526141] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526153] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 2.187 second response time
[1633526173] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526184] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.320 second response time
[1633526187] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526200] SERVICE ALERT: portalfuncionario.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1223 bytes in 0.389 second response time
[1633526212] SERVICE ALERT: rabbitmq.nasajon.com.br;[PONTO WEB] recalculo-por-apuracao-trabalhador;OK;HARD;3;RABBITMQ OK - messages OK (65) messages_ready OK (64) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633526354] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526354] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526364] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 67286 bytes in 0.174 second response time
[1633526364] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.363 second response time
[1633526564] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633526574] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.322 second response time
[1633526724] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 503 Service Temporarily Unavailable - 835 bytes in 0.314 second response time
[1633526734] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 503 Service Temporarily Unavailable - 835 bytes in 0.319 second response time
[1633526744] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 503 Service Temporarily Unavailable - 835 bytes in 0.321 second response time
[1633526755] SERVICE ALERT: multinotas.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.603 second response time
[1633527280] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (105), messages_ready OK (105) messages_unacknowledged OK (0) consumers OK (0)
[1633527291] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (114), messages_ready OK (114) messages_unacknowledged OK (0) consumers OK (0)
[1633527301] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (115), messages_ready OK (115) messages_unacknowledged OK (0) consumers OK (0)
[1633527454] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633527465] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 0.488 second response time
[1633528174] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633528185] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633528194] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633528201] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 6.186 second response time
[1633528205] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.481 second response time
[1633528225] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633528235] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.317 second response time
[1633528255] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633528275] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633528286] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.434 second response time
[1633528962] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633528975] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 3.010 second response time
[1633530585] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633530605] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633530617] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;3;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 1.533 second response time
[1633532564] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633532575] SERVICE ALERT: pontoweb.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1178 bytes in 0.384 second response time
[1633533204] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633533214] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.362 second response time
[1633536040] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (701), messages_ready OK (701) messages_unacknowledged OK (0) consumers OK (0)
[1633536051] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (702), messages_ready OK (702) messages_unacknowledged OK (0) consumers OK (0)
[1633536062] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (702), messages_ready OK (702) messages_unacknowledged OK (0) consumers OK (0)
[1633537944] SERVICE ALERT: sentry.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633537955] SERVICE ALERT: sentry.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 FOUND - 616 bytes in 0.362 second response time
[1633540376] SERVICE ALERT: lp.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633540387] SERVICE ALERT: lp.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 263 bytes in 0.528 second response time
[1633541374] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;OK;HARD;3;RABBITMQ OK - messages OK (697) messages_ready OK (0) messages_unacknowledged OK (697) consumers OK (1) All queues under the thresholds
[1633544324] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633544337] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 3.006 second response time
[1633545722] SERVICE ALERT: rabbitmq.nasajon.com.br;[MEU RH] calculo-salario;OK;HARD;3;RABBITMQ OK - messages OK (75) messages_ready OK (74) messages_unacknowledged OK (1) consumers OK (0) All queues under the thresholds
[1633546004] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633546015] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.415 second response time
[1633546564] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633546574] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 77720 bytes in 0.082 second response time
[1633548305] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633548315] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.326 second response time
[1633548424] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633548434] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 75668 bytes in 0.077 second response time
[1633548634] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633548644] SERVICE ALERT: login.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.318 second response time
[1633550064] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633550074] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 76700 bytes in 0.100 second response time
[1633550505] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633550515] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.328 second response time
[1633551656] SERVICE ALERT: v2.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: sistemabap.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: login.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: admin.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: multinotas.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: jenkins.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: atendimento.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: paineldefundo.nasajon.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551657] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551657] SERVICE ALERT: atendimento.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551657] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551657] SERVICE ALERT: diretorio.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551657] SERVICE ALERT: jenkins.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551670] SERVICE ALERT: nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: admin.nasajon.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: projetos.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: pontoweb.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: multinotas.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: diretorio.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: sentry.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: sentry.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: paineldefundo.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: lp.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: pontoweb.nasajon.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: atendimento.nasajon.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: sistemabap.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: portalfuncionario.nasajon.com.br;Check DNS;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1633551689] SERVICE ALERT: maladireta.nasajon.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;ssl_cert_expire;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551689] SERVICE ALERT: admin.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633551694] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;CRITICAL;SOFT;1;RABBITMQ CRITICAL - Connection Refused : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas : https://rabbitmq.nasajon.com.br:443/api/queues/atendimento/chamados-slas
[1633551696] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.419 second response time
[1633551696] SERVICE ALERT: atendimento.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.415 seconds response time. atendimento.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551696] SERVICE ALERT: login.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.514 seconds response time. login.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551696] SERVICE ALERT: atendimento.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 871 bytes in 0.417 second response time
[1633551696] SERVICE ALERT: jenkins.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.444 seconds response time. jenkins.nasajon.com.br returns 52.43.134.8
[1633551696] SERVICE ALERT: nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.422 seconds response time. nasajon.com.br returns 44.235.102.4,52.36.11.102,54.184.126.88,54.184.220.2
[1633551696] SERVICE ALERT: sistemabap.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.567 seconds response time. sistemabap.com.br returns 18.229.76.172,52.67.35.11,54.232.147.215
[1633551696] SERVICE ALERT: paineldefundo.nasajon.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate '*.nasajon.com.br' will expire in 323 days on 2022-08-25 23:59 +0000/UTC.
[1633551696] SERVICE ALERT: admin.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.621 seconds response time. admin.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551696] SERVICE ALERT: diretorio.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.517 seconds response time. diretorio.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551696] SERVICE ALERT: paineldefundo.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 12133 bytes in 0.531 second response time
[1633551696] SERVICE ALERT: v2.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.421 seconds response time. v2.nasajon.com.br returns 167.114.1.112
[1633551696] SERVICE ALERT: multinotas.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.487 seconds response time. multinotas.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551696] SERVICE ALERT: jenkins.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 1507288 bytes in 2.310 second response time
[1633551696] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.627 seconds response time. avaliacaodesempenho.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551699] SERVICE ALERT: pontoweb.nasajon.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate '*.nasajon.com.br' will expire in 323 days on 2022-08-25 23:59 +0000/UTC.
[1633551699] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate '*.nasajon.com.br' will expire in 323 days on 2022-08-25 23:59 +0000/UTC.
[1633551700] SERVICE ALERT: admin.nasajon.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate '*.nasajon.com.br' will expire in 323 days on 2022-08-25 23:59 +0000/UTC.
[1633551700] SERVICE ALERT: atendimento.nasajon.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate '*.nasajon.com.br' will expire in 323 days on 2022-08-25 23:59 +0000/UTC.
[1633551700] SERVICE ALERT: maladireta.nasajon.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate '*.nasajon.com.br' will expire in 323 days on 2022-08-25 23:59 +0000/UTC.
[1633551700] SERVICE ALERT: pontoweb.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.447 seconds response time. pontoweb.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551700] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.381 second response time
[1633551700] SERVICE ALERT: admin.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 845 bytes in 0.394 second response time
[1633551700] SERVICE ALERT: lp.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.384 seconds response time. lp.nasajon.com.br returns 44.235.102.4,52.36.11.102,54.184.126.88,54.184.220.2
[1633551700] SERVICE ALERT: sentry.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.361 seconds response time. sentry.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551700] SERVICE ALERT: multinotas.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.365 second response time
[1633551700] SERVICE ALERT: sentry.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 FOUND - 616 bytes in 0.368 second response time
[1633551700] SERVICE ALERT: paineldefundo.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.375 seconds response time. paineldefundo.nasajon.com.br returns 44.235.102.4,52.36.11.102,54.184.126.88,54.184.220.2
[1633551700] SERVICE ALERT: login.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.466 second response time
[1633551700] SERVICE ALERT: sistemabap.com.br;ssl_cert_expire;OK;SOFT;2;SSL OK - Certificate 'meutrabalho.app' will expire in 240 days on 2022-06-03 23:59 +0000/UTC.
[1633551700] SERVICE ALERT: projetos.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 1441 bytes in 0.381 second response time
[1633551700] SERVICE ALERT: diretorio.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 873 bytes in 0.399 second response time
[1633551700] SERVICE ALERT: portalfuncionario.nasajon.com.br;Check DNS;OK;SOFT;2;DNS OK: 0.601 seconds response time. portalfuncionario.nasajon.com.br returns 34.213.191.245,35.167.128.48,44.230.4.151,52.43.200.104
[1633551700] SERVICE ALERT: sistemabap.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 835 bytes in 0.604 second response time
[1633551700] SERVICE ALERT: nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 191184 bytes in 0.959 second response time
[1633551704] SERVICE ALERT: rabbitmq.nasajon.com.br;[ATENDIMENTO] chamados-slas;OK;SOFT;2;RABBITMQ OK - messages OK (1) messages_ready OK (0) messages_unacknowledged OK (1) consumers OK (1) All queues under the thresholds
[1633553776] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (3513), messages_ready OK (3513) messages_unacknowledged OK (0) consumers OK (0)
[1633553791] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (3570), messages_ready OK (3570) messages_unacknowledged OK (0) consumers OK (0)
[1633553802] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;WARNING;HARD;3;RABBITMQ WARNING - messages WARNING (3587), messages_ready OK (3587) messages_unacknowledged OK (0) consumers OK (0)
[1633554072] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;1;RABBITMQ WARNING - messages WARNING (3688), messages_ready OK (3686) messages_unacknowledged OK (2) consumers OK (0)
[1633554083] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;WARNING;SOFT;2;RABBITMQ WARNING - messages WARNING (4176), messages_ready OK (4174) messages_unacknowledged OK (2) consumers OK (0)
[1633554094] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-processar-documento;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (4544), messages_ready OK (4542) messages_unacknowledged OK (2) consumers OK (0)
[1633556851] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633556861] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 78172 bytes in 0.083 second response time
[1633557732] SERVICE ALERT: rabbitmq.nasajon.com.br;[MULTINOTAS] nf-schema-desconhecido;CRITICAL;HARD;3;RABBITMQ CRITICAL - messages CRITICAL (4508), messages_ready OK (4508) messages_unacknowledged OK (0) consumers OK (0)
[1633557821] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633557831] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 77397 bytes in 0.318 second response time
[1633558381] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633558391] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 76600 bytes in 0.128 second response time
[1633558820] SERVICE ALERT: maladireta.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633558830] SERVICE ALERT: maladireta.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 866 bytes in 0.336 second response time
[1633560831] SERVICE ALERT: login.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633560841] SERVICE ALERT: login.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 3730 bytes in 0.323 second response time
[1633561350] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633561361] SERVICE ALERT: avaliacaodesempenho.nasajon.com.br;url_check;OK;SOFT;2;HTTP OK: HTTP/1.1 302 Found - 899 bytes in 0.428 second response time
[1633563313] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1633563333] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1633563353] SERVICE ALERT: v2.nasajon.com.br;url_check;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1633563402] SERVICE ALERT: v2.nasajon.com.br;url_check;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 72738 bytes in 9.473 second response time